preprocessing started
[PP] input: 1. ! [X0] : mult(e,X0) = X0 [input]
[PP] input: 2. ! [X0] : e = mult(inverse(X0),X0) [input]
[PP] input: 3. ! [X0,X1,X2] : mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [input]
[PP] input: 5. ~! [X0] : mult(X0,e) = X0 [negated conjecture 4]
preprocess1 (rectify, simplify false true, flatten)
unused predicate definition removal
[PP] pred marked as built-in: =
=: +(3) -(1) 0(0)
=: +(3) -(1) 0(0)
preprocess 2 (ennf,flatten)
[PP] ennf in: 5. ~! [X0] : mult(X0,e) = X0 [negated conjecture 4]
[PP] ennf out: 6. ? [X0] : mult(X0,e) != X0 [ennf transformation 5]
naming
[PP] naming args: 1. ! [X0] : mult(e,X0) = X0 [input]
[PP] naming args: 2. ! [X0] : e = mult(inverse(X0),X0) [input]
[PP] naming args: 3. ! [X0,X1,X2] : mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [input]
[PP] naming args: 6. ? [X0] : mult(X0,e) != X0 [ennf transformation 5]
preprocess3 (nnf, flatten, skolemize)
clausify
[PP] clausify: 1. ! [X0] : mult(e,X0) = X0 [input]
[PP] clausify: 2. ! [X0] : e = mult(inverse(X0),X0) [input]
[PP] clausify: 3. ! [X0,X1,X2] : mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [input]
[PP] clausify: 8. mult(sK0,e) != sK0 [skolemisation 6,7]
[PP] final: 9. mult(e,X0) = X0 (0:5) [cnf transformation 1]
[PP] final: 10. e = mult(inverse(X0),X0) (0:6) [cnf transformation 2]
[PP] final: 11. mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) (0:11) [cnf transformation 3]
[PP] final: 12. mult(sK0,e) != sK0 (0:5) [cnf transformation 8]
preprocessing finished
[PP] onPreprocessingEnd(), Proving Helper
1. ! [X0] : mult(e,X0) = X0 [input]
2. ! [X0] : e = mult(inverse(X0),X0) [input]
3. ! [X0,X1,X2] : mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [input]
4. ! [X0] : mult(X0,e) = X0 [input]
5. ~! [X0] : mult(X0,e) = X0 [negated conjecture 4]
6. ? [X0] : mult(X0,e) != X0 [ennf transformation 5]
7. ? [X0] : mult(X0,e) != X0 => mult(sK0,e) != sK0 [choice axiom]
8. mult(sK0,e) != sK0 [skolemisation 6,7]
9. mult(e,X0) = X0 [cnf transformation 1]
10. e = mult(inverse(X0),X0) [cnf transformation 2]
11. mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [cnf transformation 3]
12. mult(sK0,e) != sK0 [cnf transformation 8]
[SA] new: 9. mult(e,X0) = X0 (0:5) [cnf transformation 1]
[SA] new: 10. e = mult(inverse(X0),X0) (0:6) [cnf transformation 2]
[SA] new: 11. mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) (0:11) [cnf transformation 3]
[SA] new: 12. mult(sK0,e) != sK0 (0:5) [cnf transformation 8]
[SA] passive: 9. mult(e,X0) = X0 (0:5) [cnf transformation 1]
[SA] passive: 10. e = mult(inverse(X0),X0) (0:6) [cnf transformation 2]
[SA] passive: 11. mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) (0:11) [cnf transformation 3]
[SA] passive: 12. mult(sK0,e) != sK0 (0:5) [cnf transformation 8]
[SA] active: 12. mult(sK0,e) != sK0 (0:5:1) [cnf transformation 8]
[SA] active: 9. mult(e,X0) = X0 (0:5:1) [cnf transformation 1]
[SA] active: 10. e = mult(inverse(X0),X0) (0:6:1) [cnf transformation 2]
[SA] active: 11. mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) (0:11:1) [cnf transformation 3]
[SA] new: 13. mult(X0,X1) = mult(e,mult(X0,X1)) (1:9) [superposition 11,9]
[SA] new: 14. mult(e,X3) = mult(inverse(X2),mult(X2,X3)) (1:10) [superposition 11,10]
[SA] new: 15. mult(mult(X4,X5),mult(X6,X7)) = mult(mult(X4,mult(X5,X6)),X7) (1:15) [superposition 11,11]
[SA] new: 16. mult(inverse(X2),mult(X2,X3)) = X3 (1:8) [forward demodulation 14,9]
[SA] passive: 16. mult(inverse(X2),mult(X2,X3)) = X3 (1:8) [forward demodulation 14,9]
[SA] new: 17. mult(X4,mult(mult(X5,X6),X7)) = mult(mult(X4,X5),mult(X6,X7)) (1:15) [forward demodulation 15,11]
[SA] new: 18. mult(X4,mult(X5,mult(X6,X7))) = mult(X4,mult(mult(X5,X6),X7)) (1:15) [forward demodulation 17,11]
[SA] active: 16. mult(inverse(X2),mult(X2,X3)) = X3 (1:8:1) [forward demodulation 14,9]
[SA] new: 19. mult(inverse(e),X0) = X0 (2:6) [superposition 16,9]
[SA] new: 20. mult(inverse(inverse(X1)),e) = X1 (2:7) [superposition 16,10]
[SA] new: 21. mult(inverse(mult(X2,X3)),mult(X2,mult(X3,X4))) = X4 (2:12) [superposition 16,11]
[SA] new: 22. mult(X5,X6) = mult(inverse(inverse(X5)),X6) (2:9) [superposition 16,16]
[SA] new: 23. mult(X1,X2) = mult(inverse(X0),mult(mult(X0,X1),X2)) (2:12) [superposition 11,16]
[SA] passive: 19. mult(inverse(e),X0) = X0 (2:6) [superposition 16,9]
[SA] passive: 20. mult(inverse(inverse(X1)),e) = X1 (2:7) [superposition 16,10]
[SA] passive: 21. mult(inverse(mult(X2,X3)),mult(X2,mult(X3,X4))) = X4 (2:12) [superposition 16,11]
[SA] passive: 22. mult(X5,X6) = mult(inverse(inverse(X5)),X6) (2:9) [superposition 16,16]
[SA] new: 24. mult(X1,X2) = mult(inverse(X0),mult(X0,mult(X1,X2))) (2:12) [forward demodulation 23,11]
[SA] active: 19. mult(inverse(e),X0) = X0 (2:6:1) [superposition 16,9]
[SA] new: 25. mult(e,X1) = X1 (3:5) [superposition 19,16]
[SA] new: 26. mult(e,X0) = X0 (3:5) [superposition 16,19]
[SA] new: 27. mult(inverse(inverse(e)),X1) = X1 (3:7) [superposition 16,19]
[SA] new: 28. mult(X2,X3) = mult(inverse(e),mult(X2,X3)) (3:10) [superposition 11,19]
[SA] passive: 27. mult(inverse(inverse(e)),X1) = X1 (3:7) [superposition 16,19]
[SA] active: 20. mult(inverse(inverse(X1)),e) = X1 (2:7:1) [superposition 16,10]
[SA] new: 29. e = mult(inverse(inverse(inverse(X0))),X0) (3:8) [superposition 16,20]
[SA] new: 30. mult(X1,X2) = mult(inverse(inverse(X1)),mult(e,X2)) (3:11) [superposition 11,20]
[SA] passive: 29. e = mult(inverse(inverse(inverse(X0))),X0) (3:8) [superposition 16,20]
[SA] new: 31. mult(X1,X2) = mult(inverse(inverse(X1)),X2) (3:9) [forward demodulation 30,9]
[SA] active: 22. mult(X5,X6) = mult(inverse(inverse(X5)),X6) (2:9:1) [superposition 16,16]
[SA] new: 32. e = mult(X0,inverse(X0)) (3:6) [superposition 22,10]
[SA] new: 33. mult(X3,e) = X3 (3:5) [superposition 22,20]
[SA] new: 34. mult(X4,mult(inverse(X4),X5)) = X5 (3:8) [superposition 22,16]
[SA] new: 35. mult(X0,e) = X0 (3:5) [superposition 20,22]
[SA] new: 36. e = mult(X1,inverse(X1)) (3:6) [superposition 10,22]
[SA] new: 37. mult(X2,mult(inverse(X2),X3)) = X3 (3:8) [superposition 16,22]
[SA] new: 38. mult(inverse(inverse(inverse(X4))),mult(X4,X5)) = X5 (3:10) [superposition 16,22]
[SA] new: 39. mult(mult(X6,X7),X8) = mult(inverse(inverse(X6)),mult(X7,X8)) (3:13) [superposition 11,22]
[SA] passive: 32. e = mult(X0,inverse(X0)) (3:6) [superposition 22,10]
[SA] passive: 33. mult(X3,e) = X3 (3:5) [superposition 22,20]
[SA] passive: 34. mult(X4,mult(inverse(X4),X5)) = X5 (3:8) [superposition 22,16]
[SA] passive: 38. mult(inverse(inverse(inverse(X4))),mult(X4,X5)) = X5 (3:10) [superposition 16,22]
[SA] new: 40. mult(X6,mult(X7,X8)) = mult(mult(X6,X7),X8) (3:11) [forward demodulation 39,22]
[SA] active: 33. mult(X3,e) = X3 (3:5:1) [superposition 22,20]
[SA] new: 41. inverse(inverse(X0)) = X0 (4:5) [superposition 33,20]
[SA] new: 42. e = inverse(e) (4:4) [superposition 33,10]
[SA] new: 43. e = inverse(e) (4:4) [superposition 33,19]
[SA] new: 44. mult(X2,e) = inverse(inverse(X2)) (4:7) [superposition 33,22]
[SA] new: 45. mult(X3,X4) = mult(X3,mult(X4,e)) (4:9) [superposition 33,11]
[SA] new: 46. e = mult(inverse(X0),X0) (4:6) [superposition 16,33]
[SA] new: 47. mult(X1,X2) = mult(X1,mult(e,X2)) (4:9) [superposition 11,33]
[SA] new: 48. mult(X3,X4) = mult(X3,mult(X4,e)) (4:9) [superposition 11,33]
[SA] new: 49. e = inverse(e) (4:4) [superposition 10,33]
[SA] new: 50. e = inverse(e) (4:4) [superposition 19,33]
[SA] new: 51. inverse(inverse(X5)) = X5 (4:5) [superposition 20,33]
[SA] new: 52. mult(X6,e) = inverse(inverse(X6)) (4:7) [superposition 22,33]
[SA] new: 53. sK0 != sK0 (4:3) [superposition 12,33]
[SA] new: 54. $false (4:0) [trivial inequality removal 53]
% Refutation found. Thanks to Tanya!
% SZS status Theorem for group
% SZS output start Proof for group
1. ! [X0] : mult(e,X0) = X0 [input]
2. ! [X0] : e = mult(inverse(X0),X0) [input]
3. ! [X0,X1,X2] : mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [input]
4. ! [X0] : mult(X0,e) = X0 [input]
5. ~! [X0] : mult(X0,e) = X0 [negated conjecture 4]
6. ? [X0] : mult(X0,e) != X0 [ennf transformation 5]
7. ? [X0] : mult(X0,e) != X0 => mult(sK0,e) != sK0 [choice axiom]
8. mult(sK0,e) != sK0 [skolemisation 6,7]
9. mult(e,X0) = X0 [cnf transformation 1]
10. e = mult(inverse(X0),X0) [cnf transformation 2]
11. mult(X0,mult(X1,X2)) = mult(mult(X0,X1),X2) [cnf transformation 3]
12. mult(sK0,e) != sK0 [cnf transformation 8]
14. mult(e,X3) = mult(inverse(X2),mult(X2,X3)) [superposition 11,10]
16. mult(inverse(X2),mult(X2,X3)) = X3 [forward demodulation 14,9]
20. mult(inverse(inverse(X1)),e) = X1 [superposition 16,10]
22. mult(X5,X6) = mult(inverse(inverse(X5)),X6) [superposition 16,16]
33. mult(X3,e) = X3 [superposition 22,20]
53. sK0 != sK0 [superposition 12,33]
54. $false [trivial inequality removal 53]
% SZS output end Proof for group
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 383
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
